---
title: 'Chapter 1 Introduction to EDA'
description: 'In this part, you will learn about the different data type, how to transform them and how to use summary statistics to get the most information out of them'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: be0e35837c
xp: 50
```

`@projector_key`
a65048197d14f669b9511b9337ad9fc9
