---
title: 'Part 1.3 Summary Statistics'
description: 'Part 1.3 Summary Statistics'
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/production/course_22747/slides/chapter5.pdf'
---

## Coursera: W2 Categorical Data: Tables, Bar Charts & Pie Charts

```yaml
type: VideoExercise
key: 243e932b66
xp: 50
```

`@projector_key`
3500671ebe3f825a4828960cd558f0cf

---

## Coursera W2: Quantitative Data: Histograms

```yaml
type: VideoExercise
key: 3d6ae5139f
xp: 50
```

`@projector_key`
0ca17f3baf5abbf27378caf18fe9d3e5

---

## Coursera W2: Quantitative Data: Numerical Summaries

```yaml
type: VideoExercise
key: 32fdd6654b
xp: 50
```

`@projector_key`
fe7491647a619c51131386b74be240aa

---

## Coursera W2 Standard Score (Empirical Rule)

```yaml
type: VideoExercise
key: 5ee15de38d
xp: 50
```

`@projector_key`
e1a87c0b68a9608aea76daedca0e71f5

---

## Coursera W2: Quantitative Data: Boxplots

```yaml
type: VideoExercise
key: c3c7e87e5f
xp: 50
```

`@projector_key`
4668db03b54efc3aff4f8afc2b139c2b

---

## Coursera W2 Demo: Interactive Histogram & Boxplot

```yaml
type: VideoExercise
key: 6d25b2d431
xp: 50
```

`@projector_key`
69a09cf1146818ffb3cff555f86e5d79

---

## Ploting Categorical Data

```yaml
type: NormalExercise
key: 09473e2a51
xp: 100
```

<!-- Guidelines for contexts: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Instruction 1
- Instruction 2

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}
data_file = 'https://assets.datacamp.com/production/repositories/5503/datasets/fa19780a7b011d9b009e8bff8e99922a8ee2eb90/housing_prices_data.csv'
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
